# Assignment 1
Before you start working on the problems, here is a small example to help you understand how to write your answers. The solution should be written within the function body given, and the final result should be returned. Then the autograder will try to call the function and validate your returned result accordingly. Before submitting to the autograder, see if the tests we have included with this file pass by pressing the *Validate* button on the toolbar.

**Note: For all assignments please write all of your code within the function we define in order to ensure it is run by the autograder correctly**

```python
def example_word_count():
    # This example question requires counting words in the example_string below.
    example_string = "Amy is 5 years old"

    # You should fix the given code only between ### FIX CODE BELOW and ### FIX CODE ABOVE. 
    # Your code should return a result. 
    # You can comment out or delete the NotImplementedError below.
    
    ### FIX CODE BELOW 
    result = example_string.split(" ")
    ### FIX CODE ABOVE
    #raise NotImplementedError()
    return len(result)
```

## Question 1: names()

Fix the incorrect regex between `### FIX CODE BELOW` and `### FIX CODE ABOVE` to generate a list of names in the simple_string.

<!-- grade_id: names, points: 0, grade: false, solution: true -->
```python
def names():
    import re
    simple_string = """Amy is 5 years old, and her sister Mary is 2 years old. 
    Ruth and Peter, their parents, have 3 kids."""

    ### FIX CODE BELOW  
    pattern = r'[A-Za-z]?'
    match = re.finditer(pattern, simple_string)
    ### FIX CODE ABOVE  
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return match
    
names()
```

<!-- grade_id: correct_names, points: 1, grade: true, solution: false -->
```python
assert len(names()) == 4, "There are four names in the simple_string"

```

## Question 2: student_grades()

The dataset file in [assets/grades.txt](assets/grades.txt) contains multiple lines of people along with their grades in a class. Fix the incorrect regex between `### FIX CODE BELOW` and `### FIX CODE ABOVE` to generate a list of just those students who received a B in the course (e.g., \[\'John Doe\', \'Jane Doe\'\].)

<!-- grade_id: grades, points: 0, grade: false, solution: true -->
```python
def student_grades():
    import re
    with open ("assets/grades.txt", "r") as file:
        grades = file.read()

    ### FIX CODE BELOW
    pattern = """(\w+)"""
    matches = re.findall(pattern,grades)
    ### FIX CODE ABOVE
        
        
    # YOUR CODE HERE
    raise NotImplementedError()

    return matches  
    
student_grades()
```

<!-- grade_id: correct_grades, points: 1, grade: true, solution: false -->
```python
assert len(student_grades()) == 16
```

## Question 3: logs()
Consider the variable 'logdata' which is a string containing a standard web log. This variable records the access a user makes when visiting a web page (like this one!). Each line of the log has the following items: 

* a host (e.g., ‘146.204.224.152’) 
* a username (e.g., ‘feest6811’ or sometime '-' since it is missing) 
* the time a request was made (e.g., ‘21/Jun/2019:15:45:24 -0700’) 
* the post request type (e.g., ‘POST /incentivize HTTP/1.1’)

Your task is to fix the incorrect regex between `### FIX CODE BELOW` and `### FIX CODE ABOVE` to convert the given data into a list of dictionaries, where each dictionary looks like the following:

```
example_dict = {"host":"146.204.224.152",
                "user_name":"feest6811",
                "time":"21/Jun/2019:15:45:24 -0700",
                "request":"POST /incentivize HTTP/1.1"}
```

<!-- grade_id: logs, points: 0, grade: false, solution: true -->
```python
def logs():
    import re
    with open("assets/logdata.txt", "r") as file:
        logdata = file.read()
    
        
    ### FIX CODE BELOW    
    pattern = """
        ([.]+)                       # host %h
        \s\-\s(.+)\s                 # user %u
        \[(.+)\]                     # time %t
        \s\"(\w\.+)\"                # request "%r"
    """
    result = [item.groupdict() for item in re.findall(pattern,logdata,re.VERBOSE)]

    ### FIX CODE ABOVE
    

    # YOUR CODE HERE
    raise NotImplementedError()

    return result

logs()
```

<!-- grade_id: correct_logs, points: 1, grade: true, solution: false -->
```python
assert len(logs()) == 979
```